<a href="https://colab.research.google.com/github/andjfog/CE880_Lab_Task/blob/main/week2/problem1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

---
# Welcome to CE880
### This is your week-2 : Problem 1 notebook
### Spring Term 2022-23
#### Prepared by Dr Haider Raza

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

Write a function to pass a list of numbers and calculate the exponential of all elements in the list format. 
`Hint`: take input as list and inside the function covert it into numpy array then perform the function. `Example` : input `[1, 4, 8]` and and output should be:  `[0.0, 2.0, 3.0]` 

<!-- grade_id: myLog_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
def myExp(mylist):
    """Write a function to pass a list of numbers and calculate the exponential of all elements """
    # YOUR CODE HERE
    input=[]
    n=np.array(input, dtype=int)
    list1= [i**2 for i in mylist]
    return list1
    raise NotImplementedError()
```

```python
in_array = [1, 2, 4, 6]
myExp(in_array)
```

Output:
```
[1, 4, 16, 36]
```

<!-- grade_id: myLog_test, points: 1, grade: true, solution: false -->
```python
"""run these test cases"""
# Test case 1
my = [1, 2, 4, 6]
assert myExp(my) == [2.718281828459045, 7.38905609893065, 54.598150033144236, 403.4287934927351]

# Test case 2
my = [1, 10, 100]
assert myExp(my) == [2.718281828459045, 22026.465794806718, 2.6881171418161356e+43]
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-15-892dc738008a> in <module>
      2 # Test case 1
      3 my = [1, 2, 4, 6]
----> 4 assert myExp(my) == [2.718281828459045, 7.38905609893065, 54.598150033144236, 403.4287934927351]
      5 
      6 # Test case 2

AssertionError:
```

---

## Question 2:

Write a function to find the k number of the smallest values from a NumPy array and remove duplicate when returning as a list. 


For example: 
`nSmallest([23, 12, 1, 3, 1, 4, 0, 4, 5, 6],3) = [0, 1, 3]`

<!-- grade_id: myfibonacci_ans, points: 0, grade: false, solution: true -->
```python
def nSmallest(array, k):
    """Find the k number of the smallest values and remove duplicates"""
    # YOUR CODE HERE
    for i in range(0, len(array)):
        if k - array[i] in array:
            return True
    return list(dict.fromkeys(array, k))
    raise NotImplementedError()
```

```python
nSmallest([100, 200, 2000, 30, 110], 4)
```

Output:
```
[100, 200, 2000, 30, 110]
```

<!-- grade_id: myfibonacci_test, points: 1, grade: true, solution: false -->
```python
"""run these test cases"""
assert nSmallest([23, 12, 1, 3, 1, 4, 0, 4, 5, 6], 3) == [0, 1, 3]
assert nSmallest([100, 200, 2000, 30, 110], 4) == [30, 100, 110, 200]
assert nSmallest([10, 10, 2, 5, 4, 5, 6, 8, 2], 2) == [2, 4]

```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-18-57578a22cadb> in <module>
      1 """run these test cases"""
----> 2 assert nSmallest([23, 12, 1, 3, 1, 4, 0, 4, 5, 6], 3) == [0, 1, 3]
      3 assert nSmallest([100, 200, 2000, 30, 110], 4) == [30, 100, 110, 200]
      4 assert nSmallest([10, 10, 2, 5, 4, 5, 6, 8, 2], 2) == [2, 4]

AssertionError:
```

---
## Question 3: 

 Write a program, which accept a list with `nan`s in it and count the maximum number of consecutive repeated `nan`.
 Example: In the `a = [nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.9999, 0.0001, 0.16, 0.101, nan, 0.16]` and answer is 2.

<!-- grade_id: max_nan_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def max_rep_nan(a):
  longest=nan
  current=nan
  for num in a:
    if num == nan:
        current +=1
    else:
        longest= max(longest, current)
        current= nan
  return max(longest, current)
  raise NotImplementedError()
```

```python
from numpy import nan
a = [nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.9999, 0.0001, 0.16, 0.101, nan, 0.16]
max_rep_nan(a)
```

Output:
```
nan
```

<!-- grade_id: max_rep_test, points: 0, grade: true, solution: false -->
```python
"""Check your answer by running this cell"""
a = [nan, nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 3

a = [nan, nan, nan, nan, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 5

a = [nan, 10, nan, 0.1, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 1

a = []
assert max_rep_nan(a) == 0
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-43-52b589a05a63> in <module>
      1 """Check your answer by running this cell"""
      2 a = [nan, nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.101, nan]
----> 3 assert max_rep_nan(a) == 3
      4 
      5 a = [nan, nan, nan, nan, nan, 0.16, 1, 0.16, 0.101, nan]

AssertionError:
```

## Pandas Data Frame 


```python
from sklearn import datasets
import pandas as pd
```

```python
# import some data to play with
diabetes = datasets.load_diabetes()
X = pd.DataFrame(diabetes.data[:, :])  # we are taking all features.
y = diabetes.target
print(X)
print(y)
```

Output:
```
0         1         2         3         4         5         6  \
0    0.038076  0.050680  0.061696  0.021872 -0.044223 -0.034821 -0.043401   
1   -0.001882 -0.044642 -0.051474 -0.026328 -0.008449 -0.019163  0.074412   
2    0.085299  0.050680  0.044451 -0.005671 -0.045599 -0.034194 -0.032356   
3   -0.089063 -0.044642 -0.011595 -0.036656  0.012191  0.024991 -0.036038   
4    0.005383 -0.044642 -0.036385  0.021872  0.003935  0.015596  0.008142   
..        ...       ...       ...       ...       ...       ...       ...   
437  0.041708  0.050680  0.019662  0.059744 -0.005697 -0.002566 -0.028674   
438 -0.005515  0.050680 -0.015906 -0.067642  0.049341  0.079165 -0.028674   
439  0.041708  0.050680 -0.015906  0.017282 -0.037344 -0.013840 -0.024993   
440 -0.045472 -0.044642  0.039062  0.001215  0.016318  0.015283 -0.028674   
441 -0.045472 -0.044642 -0.073030 -0.081414  0.083740  0.027809  0.173816   

            7         8         9  
0   -0.002592  0.019908 -0.017646  
1   -0.039493 -
```

---
## Question 4:

Write a function to find the max number in the each column in the diabetes dataset and return the output as a list. Please make sure to round the output to two decimal points. 

<!-- grade_id: filter_1_ans, points: 0, grade: false, solution: true -->
```python
def my_max_min(X):
    """Write a function to find the max number in the each column and return 
    as a list."""
    df=diabetes
    X=[]
    col= 'x'
    max_x=df.ix[df[col].idmax()]
    print(col)
    raise NotImplementedError()
```

<!-- grade_id: filter_1_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
assert my_max_min(X) == [0.11, 0.05, 0.17, 0.13, 0.15, 0.2, 0.18, 0.19, 0.13, 0.14]
```

Output:
```
<ipython-input-70-c921ad682bc9> in my_max_min(X)
      5     X=[]
      6     col= 'x'
----> 7     max_x=df.ix[df[col].idmax()]
      8     print(col)
      9     raise NotImplementedError()

/usr/local/lib/python3.8/dist-packages/sklearn/utils/__init__.py in __getattr__(self, key)
    117             return self[key]
    118         except KeyError:
--> 119             raise AttributeError(key)
    120 
    121     def __setstate__(self, state):

AttributeError: ix
```

---
## Question 5:


Write a function to calculate the `mean` and `standard deviation` of diabetes dataset (first column only). `Hint`: You can use NumPy aggregate functions. Create a list and append value to it and return the output. 

<!-- grade_id: stats_ans, points: 0, grade: false, solution: true -->
```python
def stats(X):
    """Write a functon to calculate mean and standard deviation of IRIS dataset (last column only)."""
    # YOUR CODE HERE
    list1=[]
    arr1.mean(axis=0)
    arr1.std
    append
    raise NotImplementedError()
```

```python
stats(X)
```

<!-- grade_id: stats_test, points: 2, grade: true, solution: false -->
```python
"""Check your output"""
# Test case
x = stats(X)
y= [-3.6396225400041895e-16, 0.04761904761904939]
np.testing.assert_allclose(x, y, rtol=0.05, atol=0)
```